Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **25 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Eduardo Gomez Videla"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Deterministic Quicksort

## Question 1 [time estimate: 5 minutes]

Using Figure 7.1 in Cormen et al. as a model, perform manually the partition process on the following list: A = [1,5,6,2,3,8,9,4,7]. You just need to specify the followings:
1. The array after the process is done.
2. The value of $i$ after the process is done.

<!-- grade_id: cell-7aa520f8af13679b, points: 0, grade: true, solution: true -->
Before: [1,5,6,2,3,8,9,4,7]

After: [1,5,6,2,3,4][7][8,9]

i = 4

## Question 2 [time estimate: 7 minutes]

Code up a Python implementation of `partition(A, p, r)`, closely follow the pseudo-code in Cormen et al., p.172. Your function should return the index of the pivot in the array.

<!-- grade_id: cell-1ceb2600756c60ff, points: 0, grade: false, solution: true -->
```python
def partition(A,p,r):
    """
    Assume r<len(A) and p>=0
    """
    # YOUR CODE HERE
    
    i = (p - 1)
    
    pivot = A[r] 
  
    for x in range(p, r): 
  
        if A[x] <= pivot: 
  
            # increase smallest element 
            i += 1
            A[i], A[x] = A[x], A[i] 
            
  
    A[i+1], A[r] = A[r], A[i+1] 
    return (i+1) 
    
    raise NotImplementedError()
```

Output:
```
<ipython-input-7-f0e87b302c5a> in <module>
     23     raise NotImplementedError()
     24 
---> 25 partition(1,2,3)

<ipython-input-7-f0e87b302c5a> in partition(A, p, r)
      7     i = (p - 1)
      8 
----> 9     pivot = A[r]
     10 
     11     for x in range(p, r):

TypeError: 'int' object is not subscriptable
```

<!-- grade_id: cell-a57b60117a7b82fb, points: 1, grade: true, solution: false -->
```python
A = [1,5,6,2,3,8,9,4,7]
assert(partition(A, 0, len(A)-1)==6)
```

## Question 3 [time estimate: 3 minutes]

Code up your own Python implementation of `quicksort(A, p, r)`, using `partition(A,p,r)`.

<!-- grade_id: cell-8c39ebb8cd1aa83a, points: 0, grade: false, solution: true -->
```python
def quick_sort(A,p,r):
    
    # YOUR CODE HERE
    

```

Output:
```
Sorted array is:
```

<!-- grade_id: cell-4f822430efd456ee, points: 1, grade: true, solution: false -->
```python
A = [0]
assert(quick_sort(A, 0, 0) == [0])
A = [3,1,2]
assert(quick_sort(A, 0, 2) == [1,2,3])
```

Output:
```
<ipython-input-9-3ff2a53494de> in <module>
      2 assert(quick_sort(A, 0, 0) == [0])
      3 A = [3,1,2]
----> 4 assert(quick_sort(A, 0, 2) == [1,2,3])

<ipython-input-8-a1efe82c0cd1> in quick_sort(A, p, r)
     13         # Separately sort elements before
     14         # partition and after partition
---> 15         quickSort(A, p, pi-1)
     16         quickSort(A, pi+1, r)
     17 

NameError: name 'quickSort' is not defined
```

## Question 4 [time estimate: 10 minutes] 

Explain (using experimental plots) the running time of `quick_sort` when: 
1. all elements of array A have the same value (e.g., [1,1,1])?
2. array A contains distinct elements sorted in decreasing order (e.g., [5,4,2,1])?


<!-- grade_id: cell-b58035dd5fa02329, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```