Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Eduardo Gomez Videla"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.

## Question 1 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
import math

def median(minh, maxh):
    # YOUR CODE HERE
    
    new_list = []
    
    for i in range(0,len(minh)):
        new_list.append(minh[i])
    for i in range(0,len(maxh)):
        new_list.append(maxh[i])
    
    new_list.sort()
    
    if len(new_list) % 2 == 1:

        return new_list[(math.ceil(len(new_list)/2))-1]
    
    else:
        
        middle_left = int((len(new_list)/2) -1)
        middle_right = int(middle_left+1)
        
        return (new_list[middle_left] + new_list[middle_right]) /2
        
    return(new_list)
    
    #raise NotImplementedError()
    
median([1,2,3], [3,4,5])
```

Output:
```
3.0
```

## Question 2 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


```python
# Defining some basic binary tree functions
def left(i):             # left(i): takes as input the array index of a parent node in the binary tree and 
    return 2*i + 1       #          returns the array index of its left child.

def right(i):            # right(i): takes as input the array index of a parent node in the binary tree and 
    return 2*i + 2       #           returns the array index of its right child.

def parent(i):           # parent(i): takes as input the array index of a node in the binary tree and
    return (i-1)//2      #            returns the array index of its parent



class MaxHeapq:
    """ 
    This class implements properties and methods that support a max 
    priority queue data structure
    
    Attributes
    ----------
    heap : list
        A Python list where key values in the max heap are stored
        
    heap_size : int
        An integer counter of the number of keys present in the max heap
        
    """ 
    
    
    def __init__(self):
        """
        Class initialization method.
        
        Note
        ----
        Use heapq_var = MaxHeap()
        
        """
        self.heap       = []
        self.heap_size  = 0

        
    def maxk(self):
        """
        This method returns the highest key in the priority queue
        
        Note
        ----
        Use key_var = heap_var.max()
        
        """
        return self.heap[0] 
    
     
    def heappush(self, key):   
        """
        Inserts the value of key onto the priority queue, maintaining the
        max heap invariant.
        
        Note
        ----
        Use heapq_var.heappush(key)
        
        """
        self.heap.append(-float("inf"))
        self.increase_key(self.heap_size,key)
        self.heap_size+=1
        
        
    def increase_key(self, i, key): 
        """
        This method implements the INCREASE_KEY operation, which modifies 
        the value of a key in the max priority queue with a higher value. 
        
        Note
        ----
        Use heapq_var.increase_key(i, new_key)
        
        """
        if key < self.heap[i]:
            raise ValueError('new key is smaller than the current key')
        self.heap[i] = key
        while i > 0 and self.heap[parent(i)] < self.heap[i]:
            j = parent(i)
            holder = self.heap[j]
            self.heap[j] = self.heap[i]
            self.heap[i] = holder
            i = j    
            
        
    def heapify(self, i):
        """
        This method implements the MAX_HEAPIFY operation for the max priority
        queue. The input is the array index of the root node of the subtree to 
        be heapify.
        
        Note
        ----
        Use heapq_var.heapify(i)
        
        """
        l = left(i)
        r = right(i)
        heap = self.heap
        if l <= (self.heap_size-1) and heap[l]>heap[i]:
            largest = l
        else:
            largest = i
        if r <= (self.heap_size-1) and heap[r] > heap[largest]:
            largest = r
        if largest != i:
            heap[i], heap[largest] = heap[largest], heap[i]
            self.heapify(largest)


    def heappop(self):
        """
        This method implements the EXTRACT_MAX operation. It returns the largest
        key in the max priority queue and removes this key from the max priority 
        queue.
        
        Note
        ----
        Use key_var = heapq_var.heappop() 
        
        """
        if self.heap_size < 1:
            raise ValueError('Heap underflow: There are no keys in the priority queue ')
        maxk = self.heap[0]
        self.heap[0] = self.heap[-1]
        self.heap.pop()
        self.heap_size-=1
        self.heapify(0)
        return maxk
```

<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    # YOUR CODE HERE
    
    # create new full list
    new_list = []
    
    for i in range(0,len(minh)):
        new_list.append(minh[i])
    for i in range(0,len(maxh)):
        new_list.append(maxh[i])
        
    # add new element & sort
    
    new_list.append(elem)
    
    new_list.sort()
    
    # heapify
    
    my_heap = MaxHeapq()

    for key in new_list:
        my_heap.heappush(key)
    
    return(my_heap.heap)
    
#raise NotImplementedError()

add_to_median_heap([1,2,3], [3,4,5],6)
```

Output:
```
[6, 3, 5, 1, 3, 2, 4]
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
<ipython-input-59-2a7cc854325b> in <module>
      3 for a in range(1,100,2):
      4     add_to_median_heap(minh, maxh, a)
----> 5     print(median(minh, maxh))

<ipython-input-46-a5dfdf9c293e> in median(minh, maxh)
     22         middle_right = int(middle_left+1)
     23 
---> 24         return (new_list[middle_left] + new_list[middle_right]) /2
     25 
     26     return(new_list)

IndexError: list index out of range
```

```python
add_to_median_heap([],[],1)
median()
```

Output:
```
[1]
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect(lst, k):

    # get pivot (last element) lst[-1]
    
    # Perform partition
    
    # if pivot index + 1 == k: return lst[index+1]
    
    # elif pivot index + 1 > k: recur Partition step on left list
    
    # elif pivot index + 1 < k: recur Partition step on right list


# Another non-recurring algorithm I created before realizing I was not allowed to sort
def qselect_alt(lst,k):
    # YOUR CODE HERE
    
    lst.sort()
    
    return(lst[k-1])
    
    raise NotImplementedError()
    
qselect([7,3,5,1,2],4)
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-65-253afcee666d> in <module>
      6 for a in range(100):
      7     lst2.append(qselect(lst1, a))
----> 8 assert(lst2 == list(range(100)))

AssertionError:
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))
```

Output:
```
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
<ipython-input-85-de9182726f6f> in <module>
      4 random.shuffle(lst)
      5 for a in range(100):
----> 6     print(qselect(lst, a))

<ipython-input-84-ea26b135882f> in qselect(lst, k)
     16 
     17     if k == pivot_index:
---> 18         return list[k]
     19 
     20     elif k < pivot_index:

TypeError: 'type' object is not subscriptable
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
Not recurrence relation but complexity analysis. I did this hoping it would make it easier to figure out the recurrence relation but it did not.

O(n)

E.g. k = 3, list = [1,2,5,4,3]

In this case we only move through the list once.

[1,2,3,5,4]

Now the pivot is in index 2, and index + 1 = 3. Since index + 1 = k, we return the pivot as our answer after only going through the list one time.

## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
O(n^2)

[1,2,3,4,5]